```python
def compress_str(s):
  lst = []
  sub = s[0]
  i = 1
  for x in s[1:]:
    if sub == x:
      i += 1
    else:
      lst.append((sub, i))
      sub = x
      i = 1

  lst.append((sub, i))
  return lst

```

```python
def compress_str(s):
    if not s:
        return []

    result = []
    current_char = s[0]
    count = 1

    for char in s[1:]:
        if char == current_char:
            count += 1
        else:
            result.append((current_char, count))
            current_char = char
            count = 1

    # Letztes Zeichen noch hinzufügen
    result.append((current_char, count))

    return result

```

```python
compress_str("aabccccaaa")
```

Output:
```
[('a', 2), ('b', 1), ('c', 4), ('a', 3)]
```

# Aufgabe 3.1 (3 Punkte)

Programmieren Sie eine Methode `double` in der Klasse `Point`, die die x- und y-Koordinaten des Punktes verdoppelt. Die Methode soll keinen Rückgabewert haben und die Koordinaten direkt ändern.

Beispielanwendung:
```python
>>> p = Point(3,7)
>>> p
(3|7)
>>> p.double()
>>> p
(6|14)
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class Point:
    def double(self):
        # YOUR CODE HERE
        self.x = self.x * 2
        self.y = self.y * 2


    def __init__(self, x, y):
        self.x = x
        self.y = y
    def __repr__(self):
        return "("+str(self.x)+"¦"+str(self.y)+")"
    def __eq__(self, p):
        return self.x==p.x and self.y==p.y
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
p1 = Point(3,7)
p1.double()
assert p1 == Point(6,14)
# Some hidden tests here
```

```python

```

# Aufgabe 3.3 (3 Punkte)

Programmieren Sie die `Point`-Methode `toTuple`, die einen Point in ein Tupel (`tuple`) umwandelt.

Beispielanwendung:
```python
>>> p = Point(3,7)
>>> p
(3|7)
>>> p.toTuple()
(3,7)
```

```python
class Point:
    def toTuple(self):
        # YOUR CODE HERE
        return (self.x, self.y)

    def __init__(self, x, y):
        self.x = x
        self.y = y
    def __repr__(self):
        return "("+str(self.x)+"¦"+str(self.y)+")"
    def __eq__(self, p):
        return self.x==p.x and self.y==p.y
```

```python
p1 = Point(3,7)
assert p1.toTuple() == (3,7)
# Some hidden tests here
```

# Aufgabe 3.4 (3 Punkte)

Programmieren Sie die Methode `spiegeleX` in der Klasse `Point`, die einen Punkt an der x-Achse spiegeln soll.

Hier ein Beispiel für die Verwendung der Methode:
```python
>>> p = Point(10,20)
>>> p.spiegeleX()
>>> (p.x, p.y)
(10,-20)
```

```python
class Point:
    def spiegeleX(self):
        # YOUR CODE HERE
        self.y = self.y * -1

    def __init__(self, x, y):
        self.x = x
        self.y = y
    def __repr__(self):
        return "("+str(self.x)+"¦"+str(self.y)+")"
    def __eq__(self, p):
        return self.x==p.x and self.y==p.y
```

```python
p1 = Point(10,20)
p1.spiegeleX()
assert p1 == Point(10,-20)
# Some hidden tests here
```

# Aufgabe 3.10 (2 Punkte)

Programmieren Sie eine Methode `isEmpty`, die überprüfen soll, ob eine
`Queue` leer ist. Hier ein Anwendungsbeispiel dieser Methode:

Beispielanwendungen:
```python
>>> q = Queue()
>>> q.isEmpty()
True
>>> q.enqueue('Tom')
>>> q.isEmpty()
False
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class Queue:
    def isEmpty(self):
        # Lösung hier einfügen
        # YOUR CODE HERE
        return len(self.q) == 0

    def __init__(self):
        self.q = []
    def enqueue(self,e):
        return self.q.append(e)
    def dequeue(self):
        e = self.q[0]
        self.q = self.q[1:]
        return e
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 2, grade: true, solution: false -->
```python
q = Queue()
assert q.isEmpty()
# Some hidden tests here
```

# Aufgabe 3.11 (3 Punkte)

Programmieren Sie eine Methode, die zwei Queues "addiert", indem die jeweiligen Werte als Tupel (`tuple`) zusammengefasst werden. Sollte eine Queue länger sein als die andere, dann wird der Rest einfach abgeschnitten. Stellen Sie sicher, dass die Methode eine neue `Queue` mit den erstellten Tupeln zurückgibt und nicht einfach eine Liste erstellt.

Beispielanwendung:
```python
> q1
<-- 11, 23, 45, 100, 2 <--
> q2
<-- 'Albert', 'Ayn', 'Hannah', 'Paul' <--
> q1 + q2
<-- (11, 'Albert'), (23, 'Ayn'), (45, 'Hannah'), (100, 'Paul') <--
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class Queue:

    def __add__(self, q):
        # Lösung hier einfügen
        # YOUR CODE HERE
        q3 = Queue()
        for e1, e2 in zip(self.q, q.q):
            q3.enqueue((e1, e2))
        return q3

    def isEmpty(self):
        # YOUR CODE HERE
        return len(self.q) == 0

    def __init__(self):
        self.q = []
    def enqueue(self,e):
        return self.q.append(e)
    def dequeue(self):
        e = self.q[0]
        self.q = self.q[1:]
        return e
    def __repr__(self):
        if self.isEmpty():
            return "Empty Queue"
        else:
            return "<-- " + str(self.q)[1:-1] + " <--"
    def __eq__(self,q):
        return self.q == q.q
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
q1 = Queue()
for i in [11,23,45,100,2]:
    q1.enqueue(i)
q2 = Queue()
for s in ["Albert", "Ayn", "Hannah", "Paul"]:
    q2.enqueue(s)
q3 = Queue()
for t in [(11, 'Albert'), (23, 'Ayn'), (45, 'Hannah'), (100, 'Paul')]:
    q3.enqueue(t)
assert q1+q2 == q3
# Some hidden tests here
```

# Aufgabe 3.13 (3 Punkte)

Programmieren Sie die Methode `most_common` für die `MyCounter`-Klasse; diese Methode soll eine Zahl `n` übergeben bekommen und die `n` häufigsten Elemente als Liste von Tupeln `(Anzahl, Zeichen)` zurückliefern.

Beispielanwendung:
```python
> c = MyCounter("Banana nananuna Baa")
> c.most_common(2)
[(8, 'a'), (6, 'n')]

```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class MyCounter:

    def most_common(self,n):
        # Lösung hier einfügen
        # YOUR CODE HERE
        lst = []
        for k,v in self.c.items():
            lst.append((v,k))
        return sorted(lst, reverse=True)[:n]

    def __init__(self, s=[]):
        self.c = {}
        for x in s:
            self.add(x)
    def add(self,x):
        """fuegt x in den MyCounter ein"""
        if x in self.c:
            self.c[x] = self.c[x] + 1
        else:
            self.c[x] = 1
    def __repr__(self):
        return 'MyCounter(' + str(self.c) + ')'
```

```python
c = MyCounter("Banana nananuna Baa")
c.most_common(3)
```

Output:
```
[(8, 'a'), (6, 'n'), (2, 'B')]
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
c = MyCounter("Banana nananuna Baa")
assert c.most_common(2) == [(8, 'a'), (6, 'n')]
# Some hidden tests here
```

# Aufgabe 3.14 (3 Punkte)

Offenbar ist der `+`-Operator für `MyCounter`-Objekte definiert und funktioniert wie in folgendem Beispiel gezeigt:

```python
>>> c1
MyCounter({'H': 1, 'a': 1, 'l': 2, 'o': 1})
>>> c2
MyCounter({'l': 2, 'a': 2})
>>> c1 + c2
MyCounter({'H': 1, 'a': 3, 'l': 4, 'o': 1})
```
Der `+`-Operator erzeugt also ein neues `MyCounter`-Objekt in dem die einzelnen Häufigkeiten addiert sind.
Implementieren Sie einen `+`-Operator für die `MyCounter`-Klasse der
sich gleich dem `+`-Operator der `MyCounter`-Klasse verhält.

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class MyCounter(dict):

    def __add__(self,c2):
        # Lösung hier einfügen
        # YOUR CODE HERE
        res = MyCounter()
        for k in self:
            if k in c2:
                res[k] = self[k] + c2[k]
            else:
                res[k] = self[k]
        for k in c2:
            if k not in self:
                res[k] = c2[k]
        return res

    def __init__(self, s=[]):
        super().__init__()
        for x in s:
            self.add(x)
    def add(self,x):
        """fuegt x in den MyCounter ein"""
        if x in self:
            self[x] = self[x] + 1
        else:
            self[x] = 1
    def __repr__(self):
        return 'MyCounter(' + super().__repr__() + ')'
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
assert MyCounter("Hallo")+MyCounter("lala") == MyCounter("Hallolala")
# Some hidden tests here
```

# Aufgabe 3.15 (3 Punkte)

Implementiern sie die Multiplikation für `MyCounter`-Objekte, so dass sie
wie in folgendem Beispiel funktioniert:
```python
>>> c1 = MyCounter("Hallo")
>>> c1 * 10
MyCounter({'H': 10, 'a': 10, 'l': 20, 'o': 10})
```

Die Multiplikation soll ein neues `MyCounter`-Objekt mit den Ergebnissen zurückgeben.

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class MyCounter(dict):

    def __mul__(self,n):
        # Lösung hier einfügen
        # YOUR CODE HERE
        res = MyCounter()
        if type(n) == int:
            for k in self:
                res[k] = self[k] * n
        else:
            for k in self:
                if k in m:
                    res[k] = self[k] * n[k]
                else:
                    res[k] = self[k]
            for k in m:
                if k not in self:
                    res[k] = n[k]
        return res

    def __init__(self, s=[]):
        super().__init__()
        for x in s:
            self.add(x)
    def add(self,x):
        """fuegt x in den MyCounter ein"""
        if x in self:
            self[x] = self[x] + 1
        else:
            self[x] = 1
    def __repr__(self):
        return 'MyCounter(' + super().__repr__() + ')'
```

```python
c1 = MyCounter("Hallo")
c1 * 10
```

Output:
```
MyCounter({'H': 10, 'a': 10, 'l': 20, 'o': 10})
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
c1 = MyCounter("Hallo")
assert (c1*10) == {'H': 10, 'a': 10, 'l': 20, 'o': 10}
assert (c1)  == {'H': 1, 'a': 1, 'l': 2, 'o': 1}
# Some hidden tests here
```

# Aufgabe 3.16 (3 Punkte)

Programmieren Sie die Funktion `ok(x)`, die den Parameter `x` daraufhin überprüft, ob er eine gerade ganze Zahl ist.

Beispiel:
```python
> ok(3)
False
> ok(4)
True
> ok("Hallo")
False
>ok(3.5)
False
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
def ok(x):
    # YOUR CODE HERE
    if type(x) == int:
        return x % 2 == 0
    else:
        return False
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 2, grade: true, solution: false -->
```python
assert ok(4)
assert not ok(3)
assert not ok("Hallo")
# Some hidden tests here
```

<!-- grade_id: cell-e662347fc17cfedf, points: 1, grade: true, solution: false -->
```python
# Some more hidden tests here
```

# Aufgabe 3.18a (8 Punkte)

Passen Sie die folgenden Methoden der Klasse `EvenL` so an, dass sie nur **gerade ganze Zahlen** (`int`) akzeptieren. Nutzen Sie dazu die bereits definierte Funktion `ok(x)`, um die Überprüfung durchzuführen:


- `extend` (2 Punkte)
- `insert` (2 Punkte)
- `__add__` (2 Punkte)
- `__init__` (2 Punkte)

Da `EvenL` von `list` erbt, sollen Sie die Methoden von `list` mit `super()` aufrufen.

**A**: Starten Sie mit der Anpassung von `extend`.

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class EvenL(list):

    def extend(self, it):
        # YOUR CODE HERE
        super().extend([i for i in it if ok(i)])

    def append(self, x):
        if ok(x):
            super().append(x)
    def __repr__(self):
        return "EvenL("+super().__repr__()+")"

    def ok(x):
        return type(x)==int and x%2==0
```

```python
el = EvenL([2,4])
el.extend([6,7,8,9])
```

```python
el
```

Output:
```
EvenL([2, 4, 6, 8])
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 2, grade: true, solution: false -->
```python
el = EvenL([2,4])
el.extend([6,7,8,9])
assert el == EvenL([2,4,6,8])
assert isinstance(el, list)

# Some hidden tests here
```

**B:** Fahren Sie fort mit der Anpassung von `insert`

<!-- grade_id: cell-92249e990a3beb11, points: 0, grade: false, solution: true -->
```python
class EvenL(list):

    def insert(self, i, x):
        # YOUR CODE HERE
        super().insert(i, x)

    def append(self, x):
        if ok(x):
            super().append(x)
    def __repr__(self):
        return "EvenL("+super().__repr__()+")"

def ok(x):
    return type(x)==int and x%2==0
```

<!-- grade_id: cell-be5c44951dff74d6, points: 2, grade: true, solution: false -->
```python
el = EvenL([10,4,2,7,9,7])
el.insert(4,100)
assert el == EvenL([10, 4, 2, 7, 100, 9, 7])
# Some hidden tests here
```

**C:** Fahren Sie fort mit der Anpassung von `__add__`

<!-- grade_id: cell-8e86fcb377c7dd0a, points: 0, grade: false, solution: true -->
```python
class EvenL(list):

    def __add__(self, ev):
        # YOUR CODE HERE
        #super().__add__([x for x in ev if ok(x)])
        return super().__add__([x for x in ev if ok(x)])

    def append(self, x):
        if ok(x):
            super().append(x)
    def __repr__(self):
        return "EvenL("+super().__repr__()+")"

    def ok(x):
        return type(x)==int and x%2==0
```

```python
el = [2,4,6]
el2 = el + [2,4]
el2
```

Output:
```
[2, 4, 6, 2, 4]
```

```python
el = EvenL([2,4,6])
el2 = el + [2]
print(el2)
```

Output:
```
[2, 4, 6, 2]
```

<!-- grade_id: cell-7c4d747874c63af7, points: 1, grade: true, solution: false -->
```python
el = EvenL([2,4,6])
el2 = el + [2]
assert el2 == EvenL([2,4,6,2])
# Some hidden tests here
```

<!-- grade_id: cell-5058877fc73b2e6a, points: 1, grade: true, solution: false -->
```python
# Some hidden tests here
```

**D:** Fahren Sie fort mit der Anpassung von `__init__`

<!-- grade_id: cell-61c1181a2fb8c763, points: 0, grade: false, solution: true -->
```python
class EvenL(list):

    def __init__(self, s):
        # YOUR CODE HERE
        super().__init__([x for x in s if ok(x)])

    def append(self, x):
        if ok(x):
            super().append(x)
    def __repr__(self):
        return "EvenL("+super().__repr__()+")"

def ok(x):
    return type(x)==int and x%2==0
```

```python
EvenL([1,2,3])
```

Output:
```
EvenL([2])
```

<!-- grade_id: cell-d5063934cf98d67a, points: 2, grade: true, solution: false -->
```python
assert EvenL([1,2,3]) == EvenL([2])
# Some hidden tests here
```

```python

```

# Aufgabe 3.24 (3 Punkte)

Implementieren Sie die Methode `search(k)` für die `BTree`-Klasse.  
Diese Methode soll den zum Schlüssel `k` gehörenden Wert suchen und Folgendes zurückgeben:

- Falls der Schlüssel im `BTree` enthalten ist:  
  → Ein Tupel `(Schlüssel, Wert)`  
- Falls der Schlüssel **nicht enthalten** ist:  
  → `None`

Die Methode `search` soll also folgendem Beispiel gemäss arbeiten:
```python
> bt = BTree(30,v="0")
> d = {10:"Hallo", 123: "Welt", 7: "hier", 17: "ich", 21: "und", 111: "tippe"}
> for k,v in d.items():
    bt.insert(k,v)
> bt.search(17)
(17, 'ich')
> bt.search(16) # None wird zurueckgegeben
> bt.search(30)
(30, '0')
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class BTree:
    def search(self, k):
        # YOUR CODE HERE
        if k == self.k:
            return((self.k, self.v))
        elif k > self.k:
            if self.r:
                return self.r.search(k)
            else:
                return None
        elif k < self.k:
            if self.l:
                return self.l.search(k)
            else:
                return None

    def __init__(self, k, linker_TB=None, rechter_TB=None, v=None):
        """ linker_TB muss ein BTree-Objekt sein;
            rechter_TB muss ein BTree-Objekt sein."""
        self.k = k
        self.l = linker_TB
        self.r = rechter_TB
        self.v = v
    def insert(self,k,v=None):
        if k>self.k:
            # Fuege im rechten Teilbaum ein.
            if self.r: # Es gibt einen rechten Teilbaum
                self.r.insert(k,v)
            else:      # Es gibt keinen rechten Teilbaum
                self.r = BTree(k, v=v)
        elif k<self.k:
            # Füge im linken Teilbaum ein.
            if self.l:
                self.l.insert(k,v)
            else:
                self.l = BTree(k,v=v)
        elif k==self.k:
            self.v = v
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
bt = BTree(30,v="0")
d = {10:"Hallo", 123: "Welt", 7: "hier", 17: "ich", 21: "und", 111: "tippe"}
for k,v in d.items():
    bt.insert(k,v)
assert bt.search(17) == (17, 'ich')
assert bt.search(16) == None
assert bt.search(30) == (30, '0')
# Some hidden tests here
```

Output:
```
10
17
10
17
```

# Aufgabe 3.25 (9 Punkte)

Programmieren Sie die folgenden Methoden `sumBT`, `depthBT` und `maxBT` der Klasse BTree.

## 1. `sumBT`

Programmieren Sie zunächst die Methode `sumBT`, die alle Schlüssel-Werte des
Baums aufsummiert; dazu dürfen sie annehmen, dass alle Schlüsselwerte Zahlen sind. Beispielanwendung:
```python
> bt = BTree(30)
> for k in [1,4,3]:
    bt.insert(k)
> bt.sumBT()
38
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class BTree:

    def sumBT(self):
        # YOUR CODE HERE
        summe = self.k
        print(f"Ausgangspunkt: {summe}")
        if self.r:
            summe += self.r.sumBT()
            print(f"Zwischenresultat: {summe}")
        if self.l:
            summe += self.l.sumBT()
            print(f"Zwischenresultat: {summe}")
        return summe


    def __init__(self, k, linker_TB=None, rechter_TB=None, v=None):
        """ linker_TB muss ein BTree-Objekt sein;
            rechter_TB muss ein BTree-Objekt sein."""
        self.k = k
        self.l = linker_TB
        self.r = rechter_TB
        self.v = v
    def E(self):
          # wenn es einen linken Teilbaum gibt; man kann auch schreiben: if self.l!=None
          if self.l != None:
              green_edge = [(str(self.k),str(self.l.k))]
              red_edges = self.l.E() if self.l else []
          else:
              green_edge = []
              red_edges = []
          if self.r != None:
              blue_edge = [(str(self.k),str(self.r.k))]
              black_edges = self.r.E()
          else:
              blue_edge = []
              black_edges = []
          return green_edge + red_edges + blue_edge + black_edges
    def insert(self,k,v=None):
        if k>self.k:
            # Fuege im rechten Teilbaum ein.
            if self.r: # Es gibt einen rechten Teilbaum
                self.r.insert(k,v)
            else:      # Es gibt keinen rechten Teilbaum
                self.r = BTree(k, v=v)
        elif k<self.k:
            # Füge im linken Teilbaum ein.
            if self.l:
                self.l.insert(k,v)
            else:
                self.l = BTree(k,v=v)
        elif k==self.k:
            self.v = v
```

```python
bt = BTree(10)
for k in [5,15,8,12,20,6]:
    bt.insert(k)
bt.sumBT()
```

Output:
```
Ausgangspunkt: 10
Ausgangspunkt: 15
Ausgangspunkt: 20
Zwischenresultat: 35
Ausgangspunkt: 12
Zwischenresultat: 47
Zwischenresultat: 57
Ausgangspunkt: 5
Ausgangspunkt: 8
Ausgangspunkt: 6
Zwischenresultat: 14
Zwischenresultat: 19
Zwischenresultat: 76
76
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
bt = BTree(30)
for k in [1,4,3]:
    bt.insert(k)
assert bt.sumBT() == 38
# Some hidden tests here
```

Output:
```
Ausgangspunkt: 30
Ausgangspunkt: 1
Ausgangspunkt: 4
Ausgangspunkt: 3
Zwischenresultat: 7
Zwischenresultat: 8
Zwischenresultat: 38
```

## 2. `depthBT`

Programmieren sie eine Methode `depthBT`, die die Tiefe eines Baumes
berechnet. Beispielanwendung:
```python
> bt = BTree(1)
> for k in [2,3,4]:
    bt.insert(k)
> bt.depthBT()
4
```

<!-- grade_id: cell-b2f6768c462a61df, points: 0, grade: false, solution: true -->
```python
class BTree:

    def depthBT(self):
        # YOUR CODE HERE
        if self.l:
            ldepth = self.l.depthBT()
        else:
            ldepth = 0
        if self.r:
            rdepth = self.r.depthBT()
        else:
            rdepth = 0
        return 1 + max(ldepth, rdepth)

    def __init__(self, k, linker_TB=None, rechter_TB=None, v=None):
        """ linker_TB muss ein BTree-Objekt sein;
            rechter_TB muss ein BTree-Objekt sein."""
        self.k = k
        self.l = linker_TB
        self.r = rechter_TB
        self.v = v
    def E(self):
          # wenn es einen linken Teilbaum gibt; man kann auch schreiben: if self.l!=None
          if self.l != None:
              green_edge = [(str(self.k),str(self.l.k))]
              red_edges = self.l.E() if self.l else []
          else:
              green_edge = []
              red_edges = []
          if self.r != None:
              blue_edge = [(str(self.k),str(self.r.k))]
              black_edges = self.r.E()
          else:
              blue_edge = []
              black_edges = []
          return green_edge + red_edges + blue_edge + black_edges
    def insert(self,k,v=None):
        if k>self.k:
            # Fuege im rechten Teilbaum ein.
            if self.r: # Es gibt einen rechten Teilbaum
                self.r.insert(k,v)
            else:      # Es gibt keinen rechten Teilbaum
                self.r = BTree(k, v=v)
        elif k<self.k:
            # Füge im linken Teilbaum ein.
            if self.l:
                self.l.insert(k,v)
            else:
                self.l = BTree(k,v=v)
        elif k==self.k:
            self.v = v
```

<!-- grade_id: cell-db1e270e752fa18d, points: 3, grade: true, solution: false -->
```python
bt = BTree(1)
assert bt.depthBT() == 1
for k in [2,3,4]:
    bt.insert(k)
assert bt.depthBT() == 4
# Some hidden tests here
```

## 3. `maxBT`

Programmieren Sie eine Methode `maxBT`, die das maximale Element
eines Baums zurückliefert. Beispielanwendung:
```python
> bt = BTree(1)
> for k in [2,30,4]:
    bt.insert(k)
> bt.maxBT()
30
```


<!-- grade_id: cell-cea2911c0068f559, points: 0, grade: false, solution: true -->
```python
class BTree:

    def maxBT(self):
        # YOUR CODE HERE
        if self.r:
            return self.r.maxBT()
        return self.k

    def __init__(self, k, linker_TB=None, rechter_TB=None, v=None):
        """ linker_TB muss ein BTree-Objekt sein;
            rechter_TB muss ein BTree-Objekt sein."""
        self.k = k
        self.l = linker_TB
        self.r = rechter_TB
        self.v = v
    def E(self):
          # wenn es einen linken Teilbaum gibt; man kann auch schreiben: if self.l!=None
          if self.l != None:
              green_edge = [(str(self.k),str(self.l.k))]
              red_edges = self.l.E() if self.l else []
          else:
              green_edge = []
              red_edges = []
          if self.r != None:
              blue_edge = [(str(self.k),str(self.r.k))]
              black_edges = self.r.E()
          else:
              blue_edge = []
              black_edges = []
          return green_edge + red_edges + blue_edge + black_edges
    def insert(self,k,v=None):
        if k>self.k:
            # Fuege im rechten Teilbaum ein.
            if self.r: # Es gibt einen rechten Teilbaum
                self.r.insert(k,v)
            else:      # Es gibt keinen rechten Teilbaum
                self.r = BTree(k, v=v)
        elif k<self.k:
            # Füge im linken Teilbaum ein.
            if self.l:
                self.l.insert(k,v)
            else:
                self.l = BTree(k,v=v)
        elif k==self.k:
            self.v = v
```

<!-- grade_id: cell-dd08fcf3623b5d5a, points: 3, grade: true, solution: false -->
```python
bt = BTree(1)
for k in [2,30,4]:
    bt.insert(k)
assert bt.maxBT() == 30
# Some hidden tests here
```

# Aufgabe 3.29 (3 Punkte)

**a**: Definieren Sie eine Klasse `DefaultDict`, die sich wie `dict` verhält,
ausser dass der Zugriff auf nicht definierte Schlüssel einen bei der Erstellung des `dict`-Objekts übergebenen Default-Wert zurückliefert.

**Tipp**: Dafür muss die `__getitem__`-Methode überschrieben werden.

Beispielanwendung:
```python
>>> d = DefaultDict(10)
>>> d[2]
10
>>> d[1] = 11
>>> d[1]
11
>>> d[1] + d[2] + d[3]
31
```

<!-- grade_id: cell-c6ca2726df125eb8, points: 0, grade: false, solution: true -->
```python
class DefaultDict(dict):
    def __init__(self, v):
        self.default = v
        super().__init__()
    # YOUR CODE HERE
    def __getitem__(self, key):
        if key in self:
            return super().__getitem__(key)
        else:
            return self.default
```

<!-- grade_id: cell-f0eb9a9172d178fc, points: 3, grade: true, solution: false -->
```python
d = DefaultDict(10)
assert d[2] == 10
d[1] = 11
assert d[1] == 11
assert d[1] + d[2] + d[3] == 31
# Some hidden tests here
```

**b**:
Erweitern sie nun die Klasse um eine Methode `changeDefault`, die
den Default-Wert ändert.

Beispiel:
```python
>>> d = defaultDict("x")
>>> d[1] = 10
>>> d[2]
"x"
>>> d.changeDefault("y")
>>> d[2]
"y"
```

<!-- grade_id: cell-1c06a1eff31218b8, points: 0, grade: false, solution: true -->
```python
class DefaultDict(dict):
    def __init__(self, v):
        self.default = v
        super().__init__()
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2b68f4900d3ce41a, points: 3, grade: true, solution: false -->
```python
d = DefaultDict("x")
d[1] = 10
assert d[2] == "x"
d.changeDefault("y")
assert d[2] == "y"
# Some hidden tests here
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class Point:
    def toTuple(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __init__(self, x, y):
        self.x = x
        self.y = y
    def __repr__(self):
        return "("+str(self.x)+"¦"+str(self.y)+")"
    def __eq__(self, p):
        return self.x==p.x and self.y==p.y
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
p1 = Point(3,7)
assert p1.toTuple() == (3,7)
# Some hidden tests here
```

```python

```

# Aufgabe 3.30

**7*3 = 21 Punkte**

Programmieren Sie eine Klasse Bruch, um positive Bruchzahlen zu repräsentieren.

1. Programmieren Sie den Konstruktor, der soll zwei **Attribute** speichern: 1) `z` (**Zähler**, positive ganze Zahl) und 2) `n` (**Nenner**, positive ganze Zahl)
Sie können davon ausgehen, dass sowohl der Zähler als auch der Nenner **immer positiv sind**. Eine Prüfung ist daher nicht notwendig. Siehe Beispiel unten.
2. Programmieren Sie eine String-Repräsentation von Bruch-Objekten,
so dass Ausgaben entstehen wie im Beispiel unten gezeigt.
3. Programmieren Sie die Addition zweier Brüche; sie dürfen die unten definierte Funktion `kgv` verwenden (kleinstes gemeinsames Vielfaches).
4. Programmieren Sie die Multiplikation zweier Brüche (siehe Beispiel
unten).
5. Programmieren Sie eine Methode addOne, die den Bruch um eins
erhöht – siehe Beispiel unten.
6. Programmieren Sie die Methode isNatural, die testet, ob der
Bruch eine natürliche Zahl ist.
7. Programmieren Sie die Methode reciprocal (zu deutsch: „Kehrwert“), die den Kehrwert eines Bruches zurückliefert (den Bruch
selbst jedoch unverändert lässt).
```python
>>> b = Bruch(3,4)
>>> b
3/4
>>> b2 = Bruch(1,3)
>>> b1 + b2
13/12
>>> b.addOne()
>>> b
7/4
>>> b * Bruch(1,2)
7/8
>>> Bruch(4,2).isNatural()
True
>>> b.reciprocal()
4/7
```

<!-- grade_id: cell-a3bb28a2d13ed609, points: 0, grade: false, solution: true -->
```python
import math
def kgv(a, b):
    return abs(a*b) // math.gcd(a, b)

class Bruch:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f0eb9a9172d178fc, points: 3, grade: true, solution: false -->
```python
b = Bruch(3,4)
b2 = Bruch(1,3)
b_erg = b + b2
assert (b_erg.z, b_erg.n) == (13 , 12)
b.addOne()
assert (b.z, b.n) == (7,4)
b_erg = b * Bruch(1,2)
assert (b_erg.z, b_erg.n) == (7,8)
assert Bruch(4,2).isNatural()
# Some hidden tests here
```

<!-- grade_id: cell-f6d55b21a7460a20, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 1.
```

<!-- grade_id: cell-57e46dc87ff6ab35, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 2.
```

<!-- grade_id: cell-5919bc6b58751500, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 3.
```

<!-- grade_id: cell-1239f0a41db5728a, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 4.
```

<!-- grade_id: cell-7564bbfed2b2d9e5, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 5.
```

<!-- grade_id: cell-fdd41ff980a47401, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 5.
```

<!-- grade_id: cell-a97f4b1760be7e46, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 6.
```

# Aufgabe 3.36

**7*3 = 21 Punkte**

Implementieren Sie eine Klasse für einen Aufgaben-Manager, der anstehende Aufgaben (jeweils repräsentiert als String, der die Aufgabenbeschreibung enthält) verwaltet. Jede Aufgabe hat eine Priorität (repräsentiert als eine Ganzzahl), die die Wichtigkeit oder Dringlichkeit der Aufgabe angeben soll. Je kleiner diese Zahl, desto wichtiger die entsprechende Aufgabe. Für eine Priorität kann es mehrere Aufgaben geben.  

**Datenstruktur**

Die Aufgaben werden in einem Wörterbuch (`self.aufgaben`) gespeichert.
- Der Schlüssel (`key`) ist die Priorität.
- Der Wert (`value`) ist ein Stack, in dem die Aufgaben mit dieser Priorität gespeichert werden.
- Die Klasse `Stack` ist bereits implementiert und soll für die Verwaltung der Aufgaben genutzt werden.


Der Kopf der Klasse sei gegeben durch:
```python
class AufgabenManager:
  def __init__(self):
    self.aufgaben = {}
  ...
```

Implementieren Sie die folgenden Methoden, so dass die Klasse gemäss der am Ende der Aufgabenstellung gezeigten Beispielanwendung funktioniert.
1.  Die Methode `neueAufgabe`, die eine neue Aufgabe mit einer bestimmten Priorität hinzufügt. Der Methode muss ein String und
eine Ganzzahl übergeben werden.
2. Die Methode `hoechstePrio`, die die höchste Prioritätsstufe zurückliefern soll (das heisst, die Prioritätsstufe die der niedrigsten Zahl entspricht, d.h. Prio 1 stellt z.B. eine höhrere Prio dar, als Prio 4).
3. Die Methode `erledigeNaechsteAufgabe`, die die nächste Aufgabe der höchsten Prioritätsstufe erledigt, d.h. aus der Datenstruktur löscht und den Aufgaben-String zurückliefert.
4. Eine Methode `alleAufgabenMitPrio`, die alle Aufgaben einer bestimmten Prioritätsstufe zurückliefert. Die Methode soll eine (eventuell auch leere) Liste von Strings zurückliefern.
5. Die Methode `allePrios`, die eine Liste aller Prioriätsstufen zurückliefern soll.
6. Die Methode `anzahlAufgabenPrio`, die die Anzahl der Aufgaben einer bestimmten Prioritätsstufe zurückliefern soll. Die Methode soll eine Prioritätsstufe übergeben bekommen.
7. Die Methode `anzahlAufgaben`, die zurückliefern soll, wie viele Aufgaben es insgesamt gibt.

Hier eine Beispielanwendung der Klasse:
```python
>>> aufgs = AufgabenManager()
>>> aufgs.neueAufgabe("Kueche putzen",5)
>>> aufgs.neueAufgabe("Auf Prog 1 lernen", 1)
>>> aufgs.neueAufgabe("Oma besuchen", 2)
>>> aufgs.neueAufgabe("Auf Mathe 1 lernen", 1)
>>> aufgs.neueAufgabe("Fahrrad putzen", 10)
>>> aufgs.erledigeNaechsteAufgabe()
"Auf Mathe 1 lernen"
>>> aufgs.erledigeNaechsteAufgabe()
"Auf Prog 1 lernen"
>>> print(aufgs.hoechstePrio())
2
>>> print(aufgs.anzahlAufgabenPrio(1))
0
>>> print(aufgs.anzahlAufgaben())
3
>>> print(aufgs.alleAufgabenMitPrio(10))
["Fahrrad putzen"]
>>> print(aufgs.allePrios())
[5, 2, 10]
```

<!-- grade_id: cell-a3bb28a2d13ed609, points: 0, grade: false, solution: true -->
```python
class Stack:
    def __init__(self):
        self.s = []
    def push(self,x):
        self.s.append(x)
    def pop(self):
        ret = self.s[-1]
        self.s = self.s[:-1]
        return ret
    def isEmpty(self):
        return self.s==[]
    def __repr__(self):
        return str(self.s)[1:-1]+"<->"

class AufgabenManager:
    def __init__(self):
        self.aufgaben = {}

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f63de364e6e4c723, points: 3, grade: true, solution: false -->
```python
aufgs = AufgabenManager()
aufgs.neueAufgabe("Kueche putzen",5)
aufgs.neueAufgabe("Auf Prog 1 lernen", 1)
aufgs.neueAufgabe("Oma besuchen", 2)
aufgs.neueAufgabe("Auf Mathe 1 lernen", 1)
aufgs.neueAufgabe("Fahrrad putzen", 10)
assert aufgs.erledigeNaechsteAufgabe() == 'Auf Mathe 1 lernen'
assert aufgs.erledigeNaechsteAufgabe() == 'Auf Prog 1 lernen'
assert aufgs.hoechstePrio() == 2
assert aufgs.anzahlAufgabenPrio(1) == 0
assert aufgs.anzahlAufgaben() == 3
```

<!-- grade_id: cell-1c93ff2440b40eec, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 1.
```

<!-- grade_id: cell-9326df2e250bcbc7, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 2.
```

<!-- grade_id: cell-f0eb9a9172d178fc, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 3.
```

<!-- grade_id: cell-df7112101d7f4e6f, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 4.
```

<!-- grade_id: cell-766a6cb9c68f85f8, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 5.
```

<!-- grade_id: cell-57fa709040b9aeb9, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 6.
```

<!-- grade_id: cell-0dd46fe006bf9fe8, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 7.
```

```python

```

# Aufgabe 3.37

**7*3 = 21 Punkte**

Ein **Graph** ist eine mathematische Struktur, die aus **Knoten** und **Kanten** besteht. Eine Kante kann **gerichtet oder ungerichtet** sein:  
- **Ungerichtete Graphen**: Die Kanten haben keine Richtung, d.h. eine Verbindung zwischen zwei Knoten gilt in beide Richtungen.  
- **Gerichtete Graphen**: Die Kanten haben eine Richtung, d.h. eine Verbindung von `A → B` bedeutet nicht automatisch, dass es auch `B → A` gibt.  

In dieser Übung implementieren Sie einen **gerichteten Graphen** als Adjazenzliste.  
Ein Graph kann als **Adjazenzliste** dargestellt werden - jeder Knoten hat eine Liste von Nachbarknoten (also der Knoten, mit denen er direkt verbunden ist).  

Erstellen Sie eine Klasse `Graph`, die eine Adjazenzliste als `dict` speichert:  
- Die Knoten sind ausschliesslich ganze Zahlen (Integer).  
- Eine Kante von `k` nach `j` bedeutet, dass `j` in der Liste der Nachbarn von `k` enthalten ist.  


| Methode | Beschreibung |
|---------|-------------|
| `addNode(k)` | Fügt einen neuen Knoten `k` hinzu (falls er nicht existiert). Sollte es den Knoten bereits geben, so soll einfach nichts passieren. |
| `addEdge(k, j)` | Erstellt eine Kante von `k` nach `j`, falls beide Knoten existieren und die Kante noch nicht vorhanden ist. Sollte einer der Knoten `k` oder `j` nicht existieren oder die Kante bereits existieren, so soll einfach nichts passieren. |
| `isEdge(k, j)` | Prüft, ob eine Kante von `k` nach `j` existiert (gibt `True` oder `False` zurück). |
| `V()` | Gibt eine Liste aller Knoten zurück. |
| `E()` | Gibt eine Liste aller Kanten als `(start, ziel)`-Tupel zurück. |
| `allSingles()` | Gibt eine Liste der Knoten zurück, die keine Nachbarn haben. |
| `maxEdges()` |   Gibt die maximale Anzahl an ausgehenden Kanten eines Knotens im Graphen zurück.   |
| `neighbors(v)` | Gibt eine Liste aller direkten Nachbarn von `v` zurück. |


```python
>>> g = Graph()
>>> for i in range(1,6): g.addNode(i)
>>> for i,j in [(1,2),(2,2), (1,3),
                (2,3), (3,4), (1,4),
                (4,5), (5,1)]:
       g.addEdge(i,j)
>>> g.isEdge(2,2)
True
>>> g.isEdge(3,3)
False
>>> g.V()
[1, 2, 3, 4, 5]
>>> g.E()
[(1, 2), (1, 3), (1, 4), (2, 2), (2, 3), (3, 4), (4, 5), (5, 1)]
>>> g.allSingles()
[]
>>> g.maxEdges()
3
>>> g.neighbors(1)
[2, 3, 4]



```

<!-- grade_id: cell-f63de364e6e4c723, points: 0, grade: false, solution: true -->
```python
class Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1c93ff2440b40eec, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 1.
```

<!-- grade_id: cell-e7746dcb7c45ae85, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 2.
```

<!-- grade_id: cell-9128f1f25d020095, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 3.
```

<!-- grade_id: cell-1a28a8474c9657d3, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 4.
```

<!-- grade_id: cell-84d13c60431baa8d, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 5.
```

<!-- grade_id: cell-34bc40d83d155724, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 6.
```

<!-- grade_id: cell-cc7d34cea0c5ed44, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 7.
```

<!-- grade_id: cell-a95f510b15d51a5e, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 8.
```

<!-- grade_id: cell-99fa2cda81ccca6d, points: 3, grade: true, solution: false -->
```python
# Some hidden tests for 9.
```

# Aufgabe 3.38

**4*4 = 16 Punkte**

Verwenden sie für die nachfolgenden Aufgaben aus der Bibliothek `random` ausschließlich die Funktion `randint` und/oder `sample`.
1. Schreiben Sie eine Funktion `randomMap`, die eine zufällige Permutation der Elemente einer Sequenz vornimmt und diese Permutation in
Form eines `dict`-Objekts zurückliefert.

Hinweise:
- Jedes Element der Eingabesequenz soll genau einmal als Schlüssel und genau einmal als Wert vorkommen.  
- Die Zuordnung der Elemente soll zufällig erfolgen.  
- Beispielanwendungen:
```python
>>> randomMap("ABCDEF")
        {"A":"D", "B":"B", "C":"F", "D":"E", "E":"A", "F":"C"}
>>> randomMap([10,11,12,13])
{10:10, 11:12, 12:11, 13:13}
```
2. Schreiben Sie nun eine Klasse `SubChiffre`, die einen Text mittels einer Buchstabensubstitution, die auf einer Zufallspermutation basiert, verschlüsselt. Erstellen Sie dazu in dieser Teilaufgabe zunächst die Deklaration der Klasse und den Konstruktor. Der Konstruktor
soll hierbei eine Sequenz übergeben bekommen und eine Zufallspermutation (Attribut `substMap`; siehe Beispielanwendung unten)
davon erzeugen, die für nachfolgende Verschlüsselungen und Entschlüsselungen verwendet wird.
3. Schreiben Sie die Methode `encode`, die einen "Text" bzw. Symbolfolge durch die anfangs gewählte zufällige Buchstabensubstitution verschlüsselt.
4. Schreiben Sie eine Methode `decode`, die einen Text durch die anfangs gewählte zufällige Buchstabensubstitution entschlüsselt


Hinweise:
- Bei der Verschlüsselungsmethode "Buchstabensubstitution" wird jeder Buchstabe auf genau einen anderen Buchstaben abgebildet. Welcher Buchstabe auf welchen anderen abgebildet wird, ist in einer Tabelle oder – in unserem Fall – in einem Dictionary-Objekt gespeichert. Im Fall unten wird etwa der Buchstabe `'h'` immer auf den Buchstaben `'f'` abgebildet und der Buchstabe `'l'` immer auf den Buchstaben `'o'`, usw.

Beispielanwendung der zu programmierenden Klasse:
```python
>>> g = "abcdefghijklmnopqrstuvwxyz !."
>>> c = SubChiffre(g)
>>> c.substMap
{' ': 'v',
'!': 'c',
'.': 'l',
'a': 'y',
....
>>> cipherText = c.encode("hallo welt. das ist eine nachricht!")
>>> chipherText
'fyoorvpqoscvtykvhksvqheqveydfxhdfsl'
>>> c.decode(cipherText)
```

<!-- grade_id: cell-f63de364e6e4c723, points: 0, grade: false, solution: true -->
```python
from random import randint, sample
def randomMap(seq):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1c93ff2440b40eec, points: 4, grade: true, solution: false -->
```python
# Some hidden tests for 1.
```

<!-- grade_id: cell-e6551db36791ca03, points: 0, grade: false, solution: true -->
```python
class SubChiffre:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-482e930b88dc13ca, points: 4, grade: true, solution: false -->
```python
# Some hidden tests for 2.
```

<!-- grade_id: cell-3997afbd42a28216, points: 4, grade: true, solution: false -->
```python
# Some hidden tests for 3.
```

<!-- grade_id: cell-26518db6fc36bc10, points: 4, grade: true, solution: false -->
```python
# Some hidden tests for 4.
```

<!-- grade_id: cell-dfe3f0c80000a3ab, points: 0, grade: false, solution: true -->
```python
class Point:
    def spiegeleX(self):
        # YOUR CODE HERE
        raise NotImplementedError()

    def __init__(self, x, y):
        self.x = x
        self.y = y
    def __repr__(self):
        return "("+str(self.x)+"¦"+str(self.y)+")"
    def __eq__(self, p):
        return self.x==p.x and self.y==p.y
```

<!-- grade_id: cell-fee5b1ac1e3aaa1e, points: 3, grade: true, solution: false -->
```python
p1 = Point(10,20)
p1.spiegeleX()
assert p1 == Point(10,-20)
# Some hidden tests here
```