# Musterprüfung FS24

# Teil 1 - Python Grundlagen

## Aufgabe 1
**3 Punkte**

Schreiben Sie eine Funktion mit dem Namen `list_square`, die jedes Element einer Liste in sein Quadrat verwandelt.
Die Funktion erhält eine Liste als Input und gibt eine andere Liste mit quadrierten Werten zurück.

Beispiel
```python
>>> list_square([2,3,5])
[4,9,25] 
```

<!-- grade_id: cell-792d72c221dc046e, points: 0, grade: false, solution: true -->
```python
def list_square(lst):
    #hier Ihr Code:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3a8543e0fdc701f6, points: 3, grade: true, solution: false -->
```python
#Tests
assert list_square([2,3,5]) == [4,9,25]
```

## Aufgabe 2
**3 Punkte**

Implementieren Sie die Funktion `compress_str` mit folgender Definition:

- Die Funktion nimmt einen String als Parameter an 
- Die Funktion komprimiert den String basierend auf repetierend auftretenden Zeichen
- Ausgabetyp ist eine Liste von Tupeln, bestehend aus einem Zeichen und Anzahl Wiederholungen dieses Zeichens im Originalstring.
- Falls der Input-String leer ist, soll eine leere Liste zurückgegeben werden.
  
Beispiel: 
```python
>>> compress_str("aabcccccaaa")
[("a",2), ("b", 1), ("c", 4), ("a", 3)]
```

<!-- grade_id: cell-79045e8210b3c214, points: 0, grade: false, solution: true -->
```python
def compress_str(s):
    #hier Ihr Code:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7760f15bd69a7825, points: 2, grade: true, solution: false -->
```python
# Tests
assert compress_str("aabccccaaa") == [('a', 2), ('b', 1), ('c', 4), ('a', 3)]
```

<!-- grade_id: cell-fd4d7614b9e24e69, points: 1, grade: true, solution: false -->
```python
# Tests
```

## Aufgabe 3
**3 Punkte**

Implementieren Sie die Funktion `decompress_str`, die die Umkehrfunktion von `compress_str` darstellen soll, also mit folgender Definition:

- Die Funktion nimmt Liste von Tupeln an 
- Die Funktion erzeigt daraus den passenden String und liefert diesen zurück.
 
Beispiel: 
```python
>>> decompress_str([("x", 2), ("y", 5), ("a", 1), ("b", 10)])
"xxyyyyyabbbbbbbbbb"
```

<!-- grade_id: cell-7b3303006964b32f, points: 0, grade: false, solution: true -->
```python
def decompress_str(tpl):
    #hier Ihr Code:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-586e2aa7bb1c9b08, points: 2, grade: true, solution: false -->
```python
# Tests
assert decompress_str([("x",2), ("y", 5), ("a", 1), ("b", 10)]) == "xxyyyyyabbbbbbbbbb"
```

<!-- grade_id: cell-84ce2131bd698dff, points: 1, grade: true, solution: false -->
```python
# Tests
```

## Aufgabe 4

**2 + 2 + 2 = 6 Punkte**

Wir haben eine Liste `numbers` mit Zahlen.

1. Schreiben Sie Python-Code, um alle ungeraden Zahlen aus der Liste auszuwählen und speichern Sie diese ungeraden Zahlen in der Variablen `odd_numbers`.
2.  Schreiben Sie Python-Code, um alle Zahlen aus der Liste auszuwählen, die kleiner als 50 sind und speichern Sie diese Zahlen in der Variablen `kleiner50`.
3.  Schreiben Sie Python-Code, um alle Zahlen aus der Liste auszuwählen, die mit der Ziffer `5` beginnen und speichern Sie diese Zahlen in deer Variablen `beginn5`.

<!-- grade_id: cell-4e816a79798d15f2, points: 0, grade: false, solution: true -->
```python
numbers = [701, 356, 423, 601, 830, 14, 745, 227, 379, 812, 329, 125, 314, 814, 70, 631, 
           536, 234, 453, 360, 252, 302, 893, 224, 594, 106, 783, 142, 750, 342, 746, 
           445, 505, 628, 108, 19, 642, 722, 620, 38, 74, 315, 44, 500, 351, 708, 671, 
           114, 346, 45, 205, 609, 542, 150, 476, 498, 177, 924, 651, 433, 593, 54, 364, 
           960, 326, 300, 829, 795, 818, 73]
#hier Definition von odd_numbers
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-23faac4bd798fd64, points: 2, grade: true, solution: false -->
```python
# Tests
```

<!-- grade_id: cell-6ee6fe854093d88f, points: 0, grade: false, solution: true -->
```python
numbers = [701, 356, 423, 601, 830, 14, 745, 227, 379, 812, 329, 125, 314, 814, 70, 631, 
           536, 234, 453, 360, 252, 302, 893, 224, 594, 106, 783, 142, 750, 342, 746, 
           445, 505, 628, 108, 19, 642, 722, 620, 38, 74, 315, 44, 500, 351, 708, 671, 
           114, 346, 45, 205, 609, 542, 150, 476, 498, 177, 924, 651, 433, 593, 54, 364, 
           960, 326, 300, 829, 795, 818, 73]
#hier Definition von kleiner50
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dba4cc7ed88d824d, points: 2, grade: true, solution: false -->
```python
# Tests
```

<!-- grade_id: cell-0967e780a78f966a, points: 0, grade: false, solution: true -->
```python
numbers = [701, 356, 423, 601, 830, 14, 745, 227, 379, 812, 329, 125, 314, 814, 70, 631, 
           536, 234, 453, 360, 252, 302, 893, 224, 594, 106, 783, 142, 750, 342, 746, 
           445, 505, 628, 108, 19, 642, 722, 620, 38, 74, 315, 44, 500, 351, 708, 671, 
           114, 346, 45, 205, 609, 542, 150, 476, 498, 177, 924, 651, 433, 593, 54, 364, 
           960, 326, 300, 829, 795, 818, 73]
#hier Definition von beginn5
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1fdd035a6ae825ce, points: 2, grade: true, solution: false -->
```python
# Tests
```

## Aufgabe 5
**3 Punkte**

Schreiben Sie eine Funktion namens `find_max_sum`, die eine zweidimensionale Liste als Parameter nimmt und
diejenige Liste zurückgibt, deren Elemente die grösste Summe bilden. 

Beispiel: 
```python
>>> lsts = [[1, 2, 3], [2, 3, 3], [1, 3, 3], [2, 0], [1, 2, 3, 10]]
>>> find_max_sum(lsts)
[1,2,3,10]
```

<!-- grade_id: cell-1f31319eb16377dc, points: 0, grade: false, solution: true -->
```python
def find_max_sum(lsts): 
    #Ihr Code hier
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-79026849c9535489, points: 3, grade: true, solution: false -->
```python
# Tests
assert find_max_sum([[1, 2, 3], [2, 3, 3], [1, 3, 3], [2,0], [1,2,3,10]]) == [1, 2, 3, 10]
```

## Aufgabe 6
**3 Punkte**


Schreiben Sie eine Funktion `reverseDict`, die ein dict-Objekt "umkehrt", d.h. alle Schlüssel zu Werten und alle Werte zu den entsprechenden Schlüsseln macht.
Beispiel:
```python
>>> reverseDict({1:0, 100:99, "Hallo":3, "0":"1"})
{0:1, 99:100, 3:"Hallo", "1":"0"}
```

<!-- grade_id: cell-fbf1b5c3b0bfd2fb, points: 0, grade: false, solution: true -->
```python
def reverseDict(d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-455f63bdca3fad4d, points: 3, grade: true, solution: false -->
```python
assert reverseDict({1: 0, 100: 99, "Hallo": 3, "0": "1"}) == {0: 1, 99: 100, 3: "Hallo", "1": "0"}
assert reverseDict({}) == {}
assert reverseDict({"apple": "red", "banana": "yellow", "orange": "orange"}) == {"red": "apple", "yellow": "banana", "orange": "orange"}

```

## Aufgabe 6
**3 Punkte**

Implementieren sie die Funktion `myUnion`, die eine neue Menge erzeugt,
die die Vereinigung der beiden übergebenen Mengen ist. Sie dürfen vom
Typ `set` nur die Methode add verwenden.
Beispiel:
```python
>>> myUnion({3,2,10,5}, {11,10,3,1})
{3,2,10,5,11,1}
```  

<!-- grade_id: cell-e48dcf7de30f2e05, points: 0, grade: false, solution: true -->
```python
def myUnion(set1, set2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-371781229cb17f59, points: 3, grade: true, solution: false -->
```python
assert myUnion({3, 2, 10, 5}, {11, 10, 3, 1}) == {3, 2, 10, 5, 11, 1}
assert myUnion(set(), set()) == set()
assert myUnion({3, 2, 10, 5}, set()) == {3, 2, 10, 5}
assert myUnion(set(), {11, 10, 3, 1}) == {11, 10, 3, 1}

```

# Teil 2 - Objektorientierte Programmierung

## Aufgabe 7
**3+3+3+3 = 12 Punkte**

Programmieren Sie eine Klasse `Queue`, indem Sie die folgenden Teilaufgaben lösen.

1. Programmieren Sie eine Methode `isEmpty`, die überprüfen soll, ob eine Queue leer ist. Anwendungsbeispiel siehe unten.
2. Programmieren Sie eine Methode `enqueue`, die einer Queue ein neues Element hinzufügt.
3. Programmieren Sie eine Methode `dequeue`, die das Element, das sich am längsten in der Queue befindet, herausnimmt und zurückgibt.
4. Programmieren Sie die Methode `__repr__` so, dass eine Queue die im Beispiel gezeigte String-Repräsentation besitzt. 

Beispielanwendungen:
```python
>>> q = Queue()
>>> q.isEmpty()
True
>>> q.enqueue('Tom')
>>> q.isEmpty()
False
>>> q.enqueue('Karl')
>>> q
--> Karl, Tom -->
>>> q.dequeue()
"Tom"
>>> q
--> Karl -->
```

<!-- grade_id: cell-7d97b0b356c282b7, points: 0, grade: false, solution: true -->
```python
class Queue:
    def __init__(self): 
        self.q = []
    # YOUR CODE HERE
    raise NotImplementedError()
    
    def enqueue(self,e): 
        return self.q.append(e)
    def dequeue(self):
        e = self.q[0]
        self.q = self.q[1:]
        return e
    def __repr__(self): 
        return "--> " + repr(self.q)[1:-1] + " -->"
```

<!-- grade_id: cell-8f0d2a4ff08e94a2, points: 3, grade: true, solution: false -->
```python
q = Queue()
assert q.isEmpty()
```

<!-- grade_id: cell-a28b4e75d2ff6f3f, points: 3, grade: true, solution: false -->
```python
q = Queue()
q.enqueue("Tom")
```

<!-- grade_id: cell-26d83327ba017467, points: 3, grade: true, solution: false -->
```python
q = Queue()
q.enqueue("Tom")
```

<!-- grade_id: cell-3b1eb462e1973229, points: 3, grade: true, solution: false -->
```python
q = Queue()
q.enqueue("Tom")
```

# Aufgabe 8

**3+3+3 = 9 Punkte**

Programmieren Sie eine Klasse `evenL`, die von der Klasse `list` erbt, und Listen repräsentiert, die nur gerade ganzen Zahlen enthalten können. Überschreiben Sie hierzu die folgenden `list`-Methoden so, dass ein Versuch, ein Element in eine `evenL`-liste einzufügen einfach ignoriert wird, falls es sich nicht um eine gerade ganze Zahl handelt.

1. die `append`-Methode
2. die `__add__`-Methode (die zwei Listen konkateniert)
3. die `insert`-Methode.


<!-- grade_id: cell-81cca05a1bef5ba0, points: 0, grade: false, solution: true -->
```python
def ok(x): 
    return type(x)==int and x%2==0

class EvenL(list): 
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b762b46eadac7b1d, points: 3, grade: true, solution: false -->
```python
l = EvenL([10,4,2,7,9,7])
l.insert(4,100)
assert l == EvenL([10, 4, 2, 7, 100, 9, 7])
# Some hidden tests here
```

<!-- grade_id: cell-89191cd1380a8352, points: 3, grade: true, solution: false -->
```python
l = EvenL([2,4,6])
l2 = l + [2]
assert l2 == EvenL([2,4,6,2])
# Some hidden tests here
```

<!-- grade_id: cell-adc9bef7284783ea, points: 3, grade: true, solution: false -->
```python
l = EvenL([2,4,6])
l.append(13)
assert 13 not in l
# Some hidden tests here
```

# Teil 3 - Funktionale Programmierung und Numpy

## Aufgabe 9

**2+2+2 = 6 Punkte**

1. Verwenden Sie eine List Comprehension, um eine Liste der Längen aller Methodennamen der Klasse `str` zu erzeugen; speichern Sie diese Liste in der Variablen `aufg_a`. 
2. Verwenden Sie eine List Comprehension, um die Länge des längsten Methodennamens der Klasse `str` zu berechnen; speichern Sie diese Länge in der Variablen `aufg_b`.
3. Verwenden Sie eine List Comprehension, um den Namen des längsten Methodennamen der Klasse `str` zurückzuliefern. Speichern Sie diesen Namen in der Variablen `aufg_c`

<!-- grade_id: cell-5365b067b0d5cc6b, points: 0, grade: false, solution: true -->
```python
# Hier die Variablen aufg_a, aufg_b und aufg_c definieren
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-efce5ab5564a1d83, points: 2, grade: true, solution: false -->
```python
# Some hidden tests
```

<!-- grade_id: cell-ea718e5530150f3b, points: 2, grade: true, solution: false -->
```python
# Some hidden tests
```

<!-- grade_id: cell-4e5e1175c0395225, points: 2, grade: true, solution: false -->
```python
# Some hidden tests
```

## Aufgabe 10

```python
test = """Am Anfang schuf Gott Himmel und Erde.
Und die Erde war wüst und leer, und es war finster auf der Tiefe; und der Geist Gottes schwebte auf dem Wasser.
Und Gott sprach: Es werde Licht! Und es ward Licht.
Und Gott sah, dass das Licht gut war. Da schied Gott das Licht von der Finsternis
und nannte das Licht Tag und die Finsternis Nacht. Da ward aus Abend und Morgen der erste Tag.
Da machte Gott die Feste und schied das Wasser unter der Feste von dem Wasser über der Feste. Und es geschah also.
Und Gott nannte die Feste Himmel. Da ward aus Abend und Morgen der zweite Tag.
Und Gott sprach: Es sammle sich das Wasser unter dem Himmel an besondere Orte, dass man das Trockene sehe. Und es geschah also.
Und Gott nannte das Trockene Erde, und die Sammlung der Wasser nannte er Meer. Und Gott sah, dass es gut war.
Und Gott sprach: Es lasse die Erde aufgehen Gras und Kraut, das Samen bringe, und fruchtbare Bäume, die auf der Erde Frucht bringen, in der ihr Same ist, ein jegliches nach seiner Art. Und es geschah also.
Und die Erde brachte hervor Gras und Kraut, das Samen bringt, ein jegliches nach seiner Art, und Bäume, die Frucht bringen, in der ihr Same ist, ein jegliches nach seiner Art. Und Gott sah, dass es gut war.
Da ward aus Abend und Morgen der dritte Tag.
Und Gott sprach: Es werden Lichter an der Feste des Himmels, die da scheiden Tag und Nacht und geben Zeichen, Zeiten, Tage und Jahre
und seien Lichter an der Feste des Himmels, die da leuchten auf die Erde. Und es geschah also.
Und Gott machte zwei große Lichter: das größere Licht zur Beherrschung des Tages und das kleinere Licht zur Beherrschung der Nacht und die Sterne.
Und Gott setzte sie an die Feste des Himmels, dass sie leuchten auf die Erde
und herrschen über den Tag und über die Nacht und scheiden Licht und Finsternis. Und Gott sah, dass es gut war.
Da ward aus Abend und Morgen der vierte Tag.
Und Gott sprach: Es wimmle das Wasser von lebendigem Getier, und Vögel sollen fliegen über die Erde unter der Feste des Himmels.
Und Gott schuf große Walfische und alles lebendige Getier, das sich regt, das da wimmelt in den Wassern, nach ihrer Art, und alle gefiederten Vögel nach ihrer Art. Und Gott sah, dass es gut war.
Und Gott segnete sie und sprach: Seid fruchtbar und mehret euch und füllet das Wasser in den Meeren, und die Vögel sollen sich mehren auf Erden.
Da ward aus Abend und Morgen der fünfte Tag.
Und Gott sprach: Die Erde bringe hervor lebendiges Getier, ein jedes nach seiner Art: Vieh, Gewürm und Tiere des Feldes, ein jedes nach seiner Art. Und es geschah also.
Und Gott machte die Tiere des Feldes, ein jedes nach seiner Art, und das Vieh nach seiner Art und alles Gewürm des Erdbodens nach seiner Art. Und Gott sah, dass es gut war.
Und Gott sprach: Lasset uns Menschen machen, ein Bild, das uns gleich sei, die da herrschen über die Fische im Meer und über die Vögel unter dem Himmel und über das Vieh und über die ganze Erde und über alles Gewürm, das auf Erden kriecht.
Und Gott schuf den Menschen ihm zum Bilde, zum Bilde Gottes schuf er ihn; und schuf sie als Mann und Frau.
"""
```

**2+2+2+2+2+2=12 Punkte**

1. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_1` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die Liste der jeweils ersten Worte der Zeilen zurück.
2. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_2` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die Liste aller Zeilen zurück, die mit einem Grossbuchstaben beginnen.
3. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_3` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die Länge der längsten Zeile zurück.
4. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_4` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die längsten Zeile zurück.
5. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_5` zu programmieren; diese bekommt einen Text `t` übergeben und liefert diejenige Zeile zurück, die die wenigsten Wörter enthält.
6. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_6` zu programmieren; diese bekommt einen Text `t` übergeben und liefert das längste Wort zurück.

<!-- grade_id: cell-e4af689352fbb9f6, points: 0, grade: false, solution: true -->
```python
def aufg10_1(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a71611b43743abc8, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-0cd342df14630e73, points: 0, grade: false, solution: true -->
```python
def aufg10_2(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7bb005a730fa52e7, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-ad679ca93bce36f0, points: 0, grade: false, solution: true -->
```python
def aufg10_3(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0e017c36a81c8832, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-ab49888d306763aa, points: 0, grade: false, solution: true -->
```python
def aufg10_4(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8dfef1c3dff81ea9, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-75011bff78352ec4, points: 0, grade: false, solution: true -->
```python
def aufg10_5(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5e38b4e1dcf72421, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-754848ecf6b0cb6b, points: 0, grade: false, solution: true -->
```python
def aufg10_6(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3d9294dc68ca9f35, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

## Aufgabe 11

**2+2+2 = 6 Punkte**

Gegeben sei das folgende NDArray `a`: 

```python
import numpy as np
a = np.array([256, 654, 685, 943, 278,  18, 213, 703, 410,  99, 189, 845, 498,
              604, 862, 166, 542, 356, 752, 688, 160, 331,  40, 272,  51, 416,
              174, 355, 772, 575, 186,   2, 241, 639, 780, 278, 697, 887, 357,
              105, 460,  44, 939, 849, 911,  64, 185, 516, 987, 507, 739,  30,
              594, 861, 448, 841, 172, 931, 161, 798, 135, 605, 578, 456, 657,
              8, 937, 719, 440, 106, 207, 273, 665, 477, 831, 308, 749, 734,
              701, 742])
```

Geben Sie numpy-Code an, um ...
1. alle Zahlen in `a` zurückzugeben, die kleiner als 50 sind. Speichern Sie das Ergebnis in der Variablen `a1`.
2. alle Zahlen in `a` zurückzugeben, die gerade sind. Speichern Sie das Ergebnis in der Variablen `a2`.
3. alle Zahlen in `a` auf `0` zu setzen, die grösser als 100 sind. Speichern Sie dieses Array in der Variablen `a3`.

<!-- grade_id: cell-efc385313a322698, points: 0, grade: false, solution: true -->
```python
# Definition von a1
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc9807954634cdc5, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-c1d98abd31243b0b, points: 0, grade: false, solution: true -->
```python
# Definition von a2
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4a63b4e4e76f0d26, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-5c737c4af5f11890, points: 0, grade: false, solution: true -->
```python
# Definition von a3
a3=a
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9854968229840a6c, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

```python

```

```python

```