# Nachprüfung PyF, FS24

Gesamtpunktzahl: 65 Punkte

# Teil 1 - Python Grundlagen

## Aufgabe 1
**3 Punkte**

Schreiben Sie eine Funktion `rotate_list`, die eine Liste und eine Zahl 
`n` und eine Liste als Argumente nimmt und die Liste um 
`n` Positionen rotiert. 
Beispiel:

```python
>>> rotate_list([1, 2, 3, 4, 5], 2)
[4, 5, 1, 2, 3]
```

<!-- grade_id: cell-792d72c221dc046e, points: 0, grade: false, solution: true -->
```python
def rotate_list(lst,n):
    #hier Ihr Code:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3a8543e0fdc701f6, points: 3, grade: true, solution: false -->
```python
#Tests
assert rotate_list([1, 2, 3, 4, 5], 2) == [4, 5, 1, 2, 3]
```

## Aufgabe 2
**3 Punkte** 

Schreiben Sie eine Funktion `list_difference`, die die Differenz zweier Listen berechnet, d.h. alle Elemente, die in der ersten, aber nicht in der zweiten Liste enthalten sind. 

Beispiel:
```python
>>> list_difference([1, 2, 3, 4], [2, 4, 6])
[1, 3]
```

<!-- grade_id: cell-79045e8210b3c214, points: 0, grade: false, solution: true -->
```python
def list_difference(l1,l2):
    #hier Ihr Code:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7760f15bd69a7825, points: 3, grade: true, solution: false -->
```python
# Tests
assert list_difference([1, 2, 3, 4], [2, 4, 6]) == [1, 3]
```

## Aufgabe 3

**3 Punkte**

Implementieren Sie eine Funktion `anagramm`, die testet, ob zwei Strings Anagramme voneinander sind, d.h. die gleichen Buchstaben in der gleichen Häufigkeit enthalten.

Beispiel:
```python
>>> anagramm("listen", "silent")
True
>>> anagramm("hello", "world")
False
```

<!-- grade_id: cell-7b3303006964b32f, points: 0, grade: false, solution: true -->
```python
def anagramm(s1,s2):
    #hier Ihr Code:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-586e2aa7bb1c9b08, points: 3, grade: true, solution: false -->
```python
# Tests
assert anagramm("listen", "silent")
assert not anagramm("hello", "world")
```

## Aufgabe 4

**3 Punkte**

Implementieren Sie eine Funktion `group_by_first_letter`, die eine Liste von Wörtern in einem Dictionary gruppiert, basierend auf dem ersten Buchstaben jedes Wortes. 

Beispiel:
```python
>>> group_by_first_letter(["apple", "banana", "apricot", "blueberry", "cherry"])
{"a": ["apple", "apricot"], "b": ["banana", "blueberry"], "c": ["cherry"]}
```

<!-- grade_id: cell-4e816a79798d15f2, points: 0, grade: false, solution: true -->
```python
def group_by_first_letter(words): 
    #hier Definition von odd_numbers
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-23faac4bd798fd64, points: 3, grade: true, solution: false -->
```python
# Tests
assert group_by_first_letter(["apple", "banana", "apricot", "blueberry", "cherry"]) == {'a': ['apple', 'apricot'], 'b': ['banana', 'blueberry'], 'c': ['cherry']}
```

## Aufgabe 5
**3 Punkte**

Schreiben Sie eine Funktion `remove_vowels`, die alle Vokale aus einem gegebenen String entfernt. 

Beispiel:

```python
>>> remove_vowels("hello world")
"hll wrld"
```

<!-- grade_id: cell-1f31319eb16377dc, points: 0, grade: false, solution: true -->
```python
def remove_vowels(s): 
    #Ihr Code hier
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-79026849c9535489, points: 3, grade: true, solution: false -->
```python
# Tests
assert remove_vowels("hello world") == 'hll wrld'
```

## Aufgabe 6
**3 Punkte**


Schreiben Sie eine Funktion `reverseDict`, die ein dict-Objekt "umkehrt", d.h. alle Schlüssel zu Werten und alle Werte zu den entsprechenden Schlüsseln macht.
Beispiel:
```python
>>> reverseDict({1:0, 100:99, "Hallo":3, "0":"1"})
{0:1, 99:100, 3:"Hallo", "1":"0"}
```

<!-- grade_id: cell-fbf1b5c3b0bfd2fb, points: 0, grade: false, solution: true -->
```python
def reverseDict(d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-455f63bdca3fad4d, points: 3, grade: true, solution: false -->
```python
assert reverseDict({1: 0, 100: 99, "Hallo": 3, "0": "1"}) == {0: 1, 99: 100, 3: "Hallo", "1": "0"}
assert reverseDict({}) == {}
assert reverseDict({"apple": "red", "banana": "yellow", "orange": "orange"}) == {"red": "apple", "yellow": "banana", "orange": "orange"}

```

# Teil 2 - Objektorientierte Programmierung

## Aufgabe 7

**4*3 = 12 Punkte**

Programmieren Sie eine Klasse `Bruch`, um positive Bruchzahlen zu
repräsentieren. 

Verwenden Sie den angegebenen Konstruktor, der einen ganzzahligen Zähler und einen positiven ganzzahligen Nenner benötigt. 
   
1. Programmieren Sie eine String-Repräsentation von
   `Bruch`-Objekten, so dass Ausgaben entstehen (siehe Beispiel
   unten).
2. Programmieren Sie die Multiplikation zweier Brüche (siehe
   Beispiel unten).
3. Programmieren Sie die Methode `isNatural`, die testet, ob
   der Bruch eine natürliche Zahl ist (siehe Beispiel unten).
4. Programmieren Sie die Methode `reciprocal` (zu deutsch:
   "Kehrwert"), die den Kehrwert eines Bruches zurückliefert (den
   Bruch selbst jedoch unverändert lässt; siehe Beispiel unten).

```python
>>> a = Bruch(2,5) #1. Beispiel Konstruktur
>>> b = Bruch(3,4) 
>>> b #2. String-Repräsentation
3/4
>>> a * b # 3. String-Multiplikation
6/20
>>> Bruch(4,2).isNatural() # 4. Beispiel isNatural
True
>>> b.reciprocal() # 5. Beispiel reciprocal
4/3
```


<!-- grade_id: cell-7d97b0b356c282b7, points: 0, grade: false, solution: true -->
```python
class Bruch:
    def __init__(self,z,n): 
        self.z = z
        self.n = n
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8f0d2a4ff08e94a2, points: 3, grade: true, solution: false -->
```python
a = Bruch(2,5)
b = Bruch(3,4)
assert repr(a) == "2/5"
# Test Point generation
```

<!-- grade_id: cell-7178c0cbe88a0351, points: 3, grade: true, solution: false -->
```python
assert repr(a*b) == "6/20"
```

<!-- grade_id: cell-c3530a2f94205fe7, points: 3, grade: true, solution: false -->
```python
assert Bruch(4,2).isNatural()
```

<!-- grade_id: cell-541f5a3bb2a7e9a7, points: 3, grade: true, solution: false -->
```python
assert b.reciprocal().z == 4
```

# Aufgabe 8

**3+3+3 = 9 Punkte**

Programmieren Sie eine Klasse `evenL`, die von der Klasse `list` erbt und Listen repräsentiert, die nur gerade ganze Zahlen enthalten können. Überschreiben Sie hierzu die folgenden `list`-Methoden so, dass ein Versuch, ein Element in eine `evenL`-Liste einzufügen, einfach ignoriert wird, falls es sich nicht um eine gerade ganze Zahl handelt.

1. die `append`-Methode
2. die `__add__`-Methode (die zwei Listen konkateniert)
3. die `insert`-Methode.

Sie dürfen die Funktion `ok` verwenden.

<!-- grade_id: cell-81cca05a1bef5ba0, points: 0, grade: false, solution: true -->
```python
def ok(x): 
    return type(x)==int and x%2==0

class EvenL(list): 
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b762b46eadac7b1d, points: 3, grade: true, solution: false -->
```python
l = EvenL([10,4,2,7,9,7])
l.insert(4,100)
assert l == EvenL([10, 4, 2, 7, 100, 9, 7])
# Some hidden tests here
```

<!-- grade_id: cell-89191cd1380a8352, points: 3, grade: true, solution: false -->
```python
l = EvenL([2,4,6])
l2 = l + [2]
assert l2 == EvenL([2,4,6,2])
# Some hidden tests here
```

<!-- grade_id: cell-adc9bef7284783ea, points: 3, grade: true, solution: false -->
```python
l = EvenL([2,4,6])
l.append(13)
assert 13 not in l
# Some hidden tests here
```

# Teil 3 - Funktionale Programmierung und Numpy

## Aufgabe 9

**2+2+2+2 = 8 Punkte**

1. Verwenden Sie eine List Comprehension, um eine Liste der Längen aller Methodennamen der Klasse `str` zu erzeugen, deren Namen kürzer als 7 Buchstaben ist; speichern Sie diese Liste in der Variablen `aufg_a`. 
2. Verwenden Sie eine List Comprehension, um alle Methodennamen der Klasse `str` zu erzeugen, die nicht mit einem Vokal beginnen; speichern Sie diese Liste in der Variablen `aufg_b`.
3. Verwenden Sie eine List Comprehension, um den längsten Methodennamen der Klasse `str` zu bestimmen; speichern Sie diesen in der Variablen `aufg_c`.
4. Verwenden Sie eine List Comprehension, um denjenigen Methodennamen zu bestimmen, der die meisten `i` enthält; speichern Sie diesen in der Variablen `aufg_d`. 

<!-- grade_id: cell-5365b067b0d5cc6b, points: 0, grade: false, solution: true -->
```python
# Hier die Variablen aufg_a, aufg_b, aufg_c und aufg_d definieren
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-efce5ab5564a1d83, points: 2, grade: true, solution: false -->
```python
# Some hidden tests for aufg_a
```

<!-- grade_id: cell-471bea3eb0c3c7ba, points: 2, grade: true, solution: false -->
```python
# Some hidden tests for aufg_b
```

<!-- grade_id: cell-ed897db78c208ab1, points: 2, grade: true, solution: false -->
```python
# Some hidden tests for aufg_c
```

<!-- grade_id: cell-c66a3abea9ad4dd3, points: 2, grade: true, solution: false -->
```python
# Some hidden tests for aufg_d
```

## Aufgabe 10

```python
test = """Am Anfang schuf Gott Himmel und Erde.
Und die Erde war wüst und leer, und es war finster auf der Tiefe; und der Geist Gottes schwebte auf dem Wasser.
Und Gott sprach: Es werde Licht! Und es ward Licht.
Und Gott sah, dass das Licht gut war. Da schied Gott das Licht von der Finsternis
und nannte das Licht Tag und die Finsternis Nacht. Da ward aus Abend und Morgen der erste Tag.
Da machte Gott die Feste und schied das Wasser unter der Feste von dem Wasser über der Feste. Und es geschah also.
Und Gott nannte die Feste Himmel. Da ward aus Abend und Morgen der zweite Tag.
Und Gott sprach: Es sammle sich das Wasser unter dem Himmel an besondere Orte, dass man das Trockene sehe. Und es geschah also.
Und Gott nannte das Trockene Erde, und die Sammlung der Wasser nannte er Meer. Und Gott sah, dass es gut war.
Und Gott sprach: Es lasse die Erde aufgehen Gras und Kraut, das Samen bringe, und fruchtbare Bäume, die auf der Erde Frucht bringen, in der ihr Same ist, ein jegliches nach seiner Art. Und es geschah also.
Und die Erde brachte hervor Gras und Kraut, das Samen bringt, ein jegliches nach seiner Art, und Bäume, die Frucht bringen, in der ihr Same ist, ein jegliches nach seiner Art. Und Gott sah, dass es gut war.
Da ward aus Abend und Morgen der dritte Tag.
Und Gott sprach: Es werden Lichter an der Feste des Himmels, die da scheiden Tag und Nacht und geben Zeichen, Zeiten, Tage und Jahre
und seien Lichter an der Feste des Himmels, die da leuchten auf die Erde. Und es geschah also.
Und Gott machte zwei große Lichter: das größere Licht zur Beherrschung des Tages und das kleinere Licht zur Beherrschung der Nacht und die Sterne.
Und Gott setzte sie an die Feste des Himmels, dass sie leuchten auf die Erde
und herrschen über den Tag und über die Nacht und scheiden Licht und Finsternis. Und Gott sah, dass es gut war.
Da ward aus Abend und Morgen der vierte Tag.
Und Gott sprach: Es wimmle das Wasser von lebendigem Getier, und Vögel sollen fliegen über die Erde unter der Feste des Himmels.
Und Gott schuf große Walfische und alles lebendige Getier, das sich regt, das da wimmelt in den Wassern, nach ihrer Art, und alle gefiederten Vögel nach ihrer Art. Und Gott sah, dass es gut war.
Und Gott segnete sie und sprach: Seid fruchtbar und mehret euch und füllet das Wasser in den Meeren, und die Vögel sollen sich mehren auf Erden.
Da ward aus Abend und Morgen der fünfte Tag.
Und Gott sprach: Die Erde bringe hervor lebendiges Getier, ein jedes nach seiner Art: Vieh, Gewürm und Tiere des Feldes, ein jedes nach seiner Art. Und es geschah also.
Und Gott machte die Tiere des Feldes, ein jedes nach seiner Art, und das Vieh nach seiner Art und alles Gewürm des Erdbodens nach seiner Art. Und Gott sah, dass es gut war.
Und Gott sprach: Lasset uns Menschen machen, ein Bild, das uns gleich sei, die da herrschen über die Fische im Meer und über die Vögel unter dem Himmel und über das Vieh und über die ganze Erde und über alles Gewürm, das auf Erden kriecht.
Und Gott schuf den Menschen ihm zum Bilde, zum Bilde Gottes schuf er ihn; und schuf sie als Mann und Frau.
"""
```

**4*2+4=12 Punkte**

Ein "Wort" im Sinne dieser Aufgabe ist eine durch Whitespace eingeschlossene Zeichenkette, so wie sie durch die `str`-Methode `split` zurückgegeben wird.

1. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_1` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die Liste der jeweils letzten Zeichen aller Zeilen zurück.
2. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_2` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die Liste aller Zeilen zurück, die mit einem Kleinbuchstaben beginnen.
3. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_3` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die Länge der kürzesten Zeile zurück.
4. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_4` zu programmieren; diese bekommt einen Text `t` übergeben und liefert die kürzeste Zeile zurück.
5. Verwenden Sie List Comprehensions gegebenenfalls zusammen mit anderen Funktionen, um die Funktion `aufg10_5` zu programmieren; diese bekommt einen Text `t` übergeben und liefert diejenige Zeile zurück, die das längste Wort enthält.

<!-- grade_id: cell-e4af689352fbb9f6, points: 0, grade: false, solution: true -->
```python
def aufg10_1(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a71611b43743abc8, points: 2, grade: true, solution: false -->
```python
assert aufg10_1(test) == ['.', '.', '.', 's', '.', '.', '.', '.', '.', '.', '.', '.', 'e', '.', '.', 'e', '.', '.', '.', '.', '.', '.', '.', '.', '.', '.']
```

<!-- grade_id: cell-0cd342df14630e73, points: 0, grade: false, solution: true -->
```python
def aufg10_2(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7bb005a730fa52e7, points: 2, grade: true, solution: false -->
```python
assert aufg10_2(test) == ['und nannte das Licht Tag und die Finsternis Nacht. Da ward aus Abend und Morgen der erste Tag.',
 'und seien Lichter an der Feste des Himmels, die da leuchten auf die Erde. Und es geschah also.',
 'und herrschen über den Tag und über die Nacht und scheiden Licht und Finsternis. Und Gott sah, dass es gut war.']
# Hidden tests.
```

<!-- grade_id: cell-ad679ca93bce36f0, points: 0, grade: false, solution: true -->
```python
def aufg10_3(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0e017c36a81c8832, points: 2, grade: true, solution: false -->
```python
assert aufg10_3(test) == 37
# Hidden tests.
```

<!-- grade_id: cell-ab49888d306763aa, points: 0, grade: false, solution: true -->
```python
def aufg10_4(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8dfef1c3dff81ea9, points: 2, grade: true, solution: false -->
```python
assert aufg10_4(test) == 'Am Anfang schuf Gott Himmel und Erde.'
# Hidden tests.
```

<!-- grade_id: cell-75011bff78352ec4, points: 0, grade: false, solution: true -->
```python
def aufg10_5(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5e38b4e1dcf72421, points: 4, grade: true, solution: false -->
```python
assert aufg10_5(test) == "Und Gott machte zwei große Lichter: das größere Licht zur Beherrschung des Tages und das kleinere Licht zur Beherrschung der Nacht und die Sterne." 
# Hidden tests.
```

## Aufgabe 11

**2+2+2 = 6 Punkte**

Gegeben sei das folgende NDArray `a`: 

```python
import numpy as np
a = np.array([256, 654, 685, 943, 278,  18, 213, 703, 410,  99, 189, 845, 498,
              604, 862, 166, 542, 356, 752, 688, 160, 331,  40, 272,  51, 416,
              174, 355, 772, 575, 186,   2, 241, 639, 780, 278, 697, 887, 357,
              105, 460,  44, 939, 849, 911,  64, 185, 516, 987, 507, 739,  30,
              594, 861, 448, 841, 172, 931, 161, 798, 135, 605, 578, 456, 657,
              8, 937, 719, 440, 106, 207, 273, 665, 477, 831, 308, 749, 734,
              701, 742])
```

Geben Sie numpy-Code an, um ...
1. alle Zahlen in `a` zurückzugeben, die kleiner als 100 sind. Speichern Sie das Ergebnis in der Variablen `a1`.
2. alle Zahlen in `a` zurückzugeben, die durch 19 teilbar sind. Speichern Sie das Ergebnis in der Variablen `a2`.
3. alle Zahlen in `a` auf `0` zu setzen, die grösser als 100 sind. Speichern Sie dieses Array in der Variablen `a3`.

<!-- grade_id: cell-efc385313a322698, points: 0, grade: false, solution: true -->
```python
# Definition von a1
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc9807954634cdc5, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-c1d98abd31243b0b, points: 0, grade: false, solution: true -->
```python
# Definition von a2
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4a63b4e4e76f0d26, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```

<!-- grade_id: cell-5c737c4af5f11890, points: 0, grade: false, solution: true -->
```python
# Definition von a3
a3=a
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9854968229840a6c, points: 2, grade: true, solution: false -->
```python
# Hidden tests.
```